# Worksheet 6: Introduction to noSQL databases and document model with MongoDB

## Exercise 1: Recap concepts in non-relational database

{rubric: accuracy = 5}




1. **What is a key difference between relational and non-relational databases?**
   - A) Relational databases use a flexible schema, while non-relational databases use a fixed schema.
   - B) Relational databases use SQL for queries, while non-relational databases use various query languages.
   - C) Non-relational databases are always faster than relational databases.
   - D) Relational databases do not support ACID properties, while non-relational databases do.


My Answer is B
> YOUR ANSWER HERE


2. **Which of the following is a disadvantage of using a relational database?**
   - A) They are not suitable for complex queries.
   - B) They require a fixed schema.
   - C) They do not support transactions.
   - D) They cannot handle large volumes of data.
  


  > YOUR ANSWER HERE
  My answer is still B cause other options like ACD,relational data base can also do that

3. **Which type of non-relational database is best suited for hierarchical & networked data storage?**
   - A) Column-based
   - B) Key-value
   - C) Graph
   - D) Document


My answer is C
> YOUR ANSWER HERE

If this doesn't work, you should install `pymongo` in your conda environment by typing `!pip install pymongo` in a new code cell

In this exercise, we will create a free cluster on mongoDB Atlas and load sample datasets. 

### 2.1 Set up your MongoDB Atlas cluster

{rubric: completion = 5}

Please follow the steps below to create your free cluster on MongoDB Atlas:

1. Navigate to https://www.mongodb.com/cloud/atlas/register and sign up a free account with your email
2. Open your email and verify the account
3. Log in to your MongoDB account and create a free cluster, make sure to have the following configurations:
    - Check the M0 tier
    - Uncheck the 'Preload sample dataset' option
![](img/mongodb1.png)


4. Click `Create Deployment`
5. It's going to pop up a screen to Connect to Cluster():
    - Now you will need to add the first user to this database, choose an `username` and `password` that you can easily remember. 
    - Click choose a connection method
![](img/mongodb2.png)
    - Choose Drivers
    - Note down your host information. It should look something like `cluster0.xxxx.mongodb.net`
![](img/mongodb3.png)
6. Now open your `credientials.json` file in worksheet6 repo and fill in the username, password, and host information. Save the file.
7. Once the cluster has been created:
   - Click `Database` on the left tab
   - Click the triple dot `... button`
   - Select `Load Sample Dataset`
![](img/mongodb4.png)

### 2.2 

{rubric: accuracy = 5}

After loading the data, open the `sample_restaurant` database and answer the following questions:

- How many collections are present in the `sample_restaurant` database?
- List the number of documents and the storage size of each collection in the `sample_restaurant` database


Hints: You can read more about the documentation of `sample_restaurant` database here https://www.mongodb.com/docs/atlas/sample-data/sample-restaurants/

### 2.3

{rubric: accuracy = 5}

Open the `restaurants` collection. Filter for this ObjectId 

> {_id: ObjectId('5eb3d668b31de5d588f4292c')}

Answer following questions:

- What are the coordinates of that restaurant?
It is Array (2)
- Which borough does this restaurant locate in?
It is "Staten Island"
- How many reviews did this restaurant get?
There is 45 reviews
- What is the lowest rating score?
It is 9
- What is the highest rating score?
It is 12

## Exercise 3: Connect to your mongoDB using pymongo

What you need:
- The host URL of your mongoDB connection (should look something like `cluster0.lqirl.mongodb.net`)
- Your mongoDB username
- Your mongoDB password

Modify the `credentials.json` file with appropriate information above.

Below is the starter code to connect to your MongoDB database, provided that you have a `credentials.json` file with the correct information.

```python
from pymongo import MongoClient # import mongo client to connect
import json # import json to load credentials
import urllib.parse

# load credentials from json file
with open('credentials.json') as f:
    login = json.load(f)

# assign credentials to variables
username = login['username']
password = urllib.parse.quote(login['password'])
host = login['host']
url = "mongodb+srv://{}:{}@{}/?retryWrites=true&w=majority".format(username, password, host)
```

```python
# connect to the database
client = MongoClient(url)
```

### 3.1 

{rubric: accuracy = 5}

Write code to list all databases in the client server

Hint: See the lecture 12 for example

```python
# list all databases
for db_info in client.list_database_names():
   print(db_info)
```

Output:
```
sample_airbnb
sample_geospatial
sample_mflix
sample_restaurants
sample_supplies
sample_weatherdata
admin
local
```

### 3.2 

{rubric: accuracy = 5}

Write code to list all collections in the sample_restaurants database

Hint: See the lecture 12 for example

```python
db = client['sample_restaurants']

# List all collections
collections = db.list_collection_names()
for collection in collections:
    # Print the collections
    print(collection)

```

Output:
```
restaurants
neighborhoods
```

### 3.3 

{rubric: accuracy = 5}

Display the first document in the restaurants collection in the sample_restaurants database

Hint: See the lecture 12 for example

```python
# show the first document
db = client['sample_restaurants']

# Access the restaurants collection
restaurants_collection = db['restaurants']

# Retrieve and display the first document in the collection
first_document = restaurants_collection.find_one()

# Print the first document
print(first_document)
```

Output:
```
{'_id': ObjectId('5eb3d668b31de5d588f4292a'), 'address': {'building': '2780', 'coord': [-73.98241999999999, 40.579505], 'street': 'Stillwell Avenue', 'zipcode': '11224'}, 'borough': 'Brooklyn', 'cuisine': 'American', 'grades': [{'date': datetime.datetime(2014, 6, 10, 0, 0), 'grade': 'A', 'score': 5}, {'date': datetime.datetime(2013, 6, 5, 0, 0), 'grade': 'A', 'score': 7}, {'date': datetime.datetime(2012, 4, 13, 0, 0), 'grade': 'A', 'score': 12}, {'date': datetime.datetime(2011, 10, 12, 0, 0), 'grade': 'A', 'score': 12}], 'name': 'Riviera Caterer', 'restaurant_id': '40356018'}
```

## Submission instructions

{rubric: mechanics = 5}

- Make sure the notebook can run from top to bottom without any error. Restart the kernel and run all cells.
- Commit and push your notebook to the github repo
- Double check your notebook is rendered properly on Github and you can see all the outputs clearly

Tests for q1 (points: 1)
```
>>> assert sum(1, 2) == 3, 'Your function is not implemented correctly'
```
